```python
# Initialize Otter
import otter
grader = otter.Notebook("DSA2024Test.ipynb")
```

## DSA 2024 Summer School Admittance Check

Thanks for your interest in attending DSA 2024 Nyeri, Kenya. To attend the summer school you have to have some level of basic Python proficiency. Completing the following notebook should ensure you have the right kind of background to benefit maximally from the Summer School. See you in Nyeri!

```python
import pandas as pd
import numpy as np
%matplotlib inline
import otter

grader = otter.Notebook()
```

**Question 1:** write a function `isValid(s)` that takes as argument a string s containing a sequence of parenthesis '(', ')', '{', '}', '[' and ']', and  determines if the input is valid. A input string is valid if for every open parenthensis there is a close one and parenthesis is well-formed. e.g  "(){}[]" is valid.

```python
# def isValid(s: str) -> bool:
def isValid(s: str) -> bool:
    stack = []
    mapping = {')': '(', '}': '{', ']': '['}

    for character in s:
        if character in mapping.values():
            stack.append(character)
        elif character in mapping.keys():
            if not stack or mapping[character] != stack.pop():
                return False
        else:
            # Ignore non-parenthesis characters
            continue

    return len(stack) == 0

# Test cases
print(isValid("(){}[]"))  # True
print(isValid("([{}])"))  # True
print(isValid("(]"))      # False
print(isValid("([)]"))    # False
print(isValid("{[]}"))    # True

```

Output:
```
True
True
False
False
True
```

```python
grader.check("q1")
```

Output:
```
q1 results: All test cases passed!
```

**Question 2:** Given a paragraph as a string, write a function that return the number of character with odd frequencies. E.g The paragraph ``DSA 2024 Nyeri`` has *10* characters with odd frequencies. i.e the entire frequency count is given as {' ': 2, '2': 2, 'D': 1, 'S': 1, 'A': 1, '0': 1, '4': 1, 'N': 1, 'y': 1, 'e': 1, 'r': 1, 'i': 1}) and there are *10* characters with odd frequences. So the function should return *10*. 

```python
def oddFrequencyCounter(theParagraph: str) -> int:
    # Dictionary to store the frequency of each character
    frequency_count = {}
    
    # Count the frequency of each character
    for char in theParagraph:
        frequency_count[char] = frequency_count.get(char, 0) + 1
    
    # Count the number of characters with odd frequencies
    odd_frequency_count = sum(1 for count in frequency_count.values() if count % 2 != 0)
    
    return odd_frequency_count

# Example usage
paragraph = "DSA 2024 Nyeri"
result = oddFrequencyCounter(paragraph)
print(f"The paragraph '{paragraph}' has {result} characters with odd frequencies.")

                         
#     ...
```

Output:
```
The paragraph 'DSA 2024 Nyeri' has 10 characters with odd frequencies.
```

```python
grader.check("q2")
```

Output:
```
q2 results: All test cases passed!
```

**Question 3:** Write an infinite generator function `odd_squares_sum` that yields the sum of square of odd numbers. e.g $1^2 + 3^2 + 5^2 + ...$ up to a ``limit``

```python
def odd_squares_sum(limit):
    
    """
    Generator function that yields the sum of squares of odd numbers up to a limit.

    Args:
        limit: The upper limit for the sum.

    Returns:
        The final sum of squares of odd numbers up to the limit.
    """

    sum_value = 0
    n = 1

    while sum_value < limit:
        sum_value += n ** 2
        n += 2

    return sum_value

final_sum = odd_squares_sum(20)
print(final_sum)
```

Output:
```
35
```

```python
grader.check("q3")
```

Output:
```
q3 results:
    q3 - 1 result:
        ❌ Test case failed
        Trying:
            points: 4
        Expecting nothing
        ok
        Trying:
            theGen = odd_squares_sum(12)
        Expecting nothing
        ok
        Trying:
            assert next(theGen) == 1
        Expecting nothing
        **********************************************************************
        Line 3, in q3 0
        Failed example:
            assert next(theGen) == 1
        Exception raised:
            Traceback (most recent call last):
              File "/home/kipkoech/anaconda3/lib/python3.11/doctest.py", line 1351, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q3 0[2]>", line 1, in <module>
                assert next(theGen) == 1
                       ^^^^^^^^^^^^
            TypeError: 'int' object is not an iterator
        Trying:
            assert next(theGen) == 9
        Expecting nothing
        ********************
```

**Question 4:** Using the `odd_squares_sum` generator defined above, create a list of sum of squares up to a limit of $20$ and store the results in a numpy.array variable called `oddSumList` 

```python
for x in myvalue:
    mylist= [ mylist.append(x)]
    mylist
```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In[33], line 1
----> 1 for x in myvalue:
      2     mylist= [ mylist.append(x)]
      3     mylist

TypeError: 'int' object is not iterable
```

```python
grader.check("q4")
```

**Question 5:** Compute the element-wise remainder of ``oddSumList`` when divided by $5$ and merge it with ``oddSumList``. The final output stored in the variable `mergedList` should be in the form of a list of tupples e.g ``[(1,1), (4,9), (0,25), ...]`` 

```python
import numpy as np

# Create the list of odd numbers
odd_numbers = np.arange(1, 100, 2)

# Calculate the sum of odd numbers up to and including the last element
odd_sum_list = np.cumsum(odd_numbers)

# Calculate the element-wise remainder when divided by 5
remainders = odd_sum_list % 5

# Merge the odd sum list and remainders into a list of tuples
merged_list = list(zip(odd_sum_list, remainders))

# Print the merged list
```

```python
grader.check("q5")
```

Output:
```
q5 results:
    q5 - 1 result:
        ❌ Test case failed
        Trying:
            points: 2
        Expecting nothing
        ok
        Trying:
            assert mergedList == [(1, 1), (4, 9), (0, 25), (4, 49), (1, 81), (1, 121), (4, 169), (0, 225), (4, 289), (1, 361)]
        Expecting nothing
        **********************************************************************
        Line 2, in q5 0
        Failed example:
            assert mergedList == [(1, 1), (4, 9), (0, 25), (4, 49), (1, 81), (1, 121), (4, 169), (0, 225), (4, 289), (1, 361)]
        Exception raised:
            Traceback (most recent call last):
              File "/home/kipkoech/anaconda3/lib/python3.11/doctest.py", line 1351, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q5 0[1]>", line 1, in <module>
                assert mergedList == [(1, 1), (4, 9), (0, 25), (4, 49), (1, 81), (1, 121), (4, 169), (0, 225), (4, 289), (1, 361)]
                      
```

**Question 6:**  Write a function `greatest_common_divisor` that takes two inputs `a` and `b` and returns the greatest common divisor of the two numbers. E.g. input `(10, 15)` would return `5`

```python
def greatest_common_divisor(a, b):
    a= abs(a)
    b= abs(b)
    while b!=0:
        a, b = b, a % b
    return a
print(greatest_common_divisor(10,15))
```

Output:
```
5
```

```python
grader.check("q6")
```

Output:
```
q6 results: All test cases passed!
```

**Question 7:**  Write a function `get_3_nearest` that takes in a point of interest ``pt`` and a **list** of points ``ptlist``  and returns a list of 3 nearest points from the point of interest ``pt``. Assume the distance between any two point is defined by the `L1-norm`.

```python

def get_3_nearest(pt, ptlist):
  """
  This function finds the 3 nearest points in a list of points to a given point of interest.

  Args:
    pt: The point of interest.
    ptlist: A list of points.

  Returns:
    A list of the 3 nearest points in ptlist to pt.
  """

  # Ensure all points have the same dimensionality
  if len(pt) != len(ptlist[0]):
    raise ValueError("All points must have the same dimensionality")

  # Calculate the distances between the point of interest and all other points
  distances = [L1_norm(pt, p) for p in ptlist]

  # Sort the distances and select the indices of the 3 smallest distances
  nearest_indices = np.argsort(distances)[:3]

  # Return the 3 nearest points using the indices
  return [ptlist[i] for i in nearest_indices]

# Define the L1-norm function
def L1_norm(pt1, pt2):
    return np.sum(np.abs(pt1 - pt2))

pt = np.array([1, 2])
ptlist = [np.array([3, 4]), np.array([5, 6]), np.array([7, 8]), np.array([9, 10])]

nearest_points = get_3_nearest(pt, ptlist)

print(f"3 nearest points to {pt}: {nearest_points}")
    
```

Output:
```
3 nearest points to [1 2]: [array([3, 4]), array([5, 6]), array([7, 8])]
```

```python
grader.check("q7")
```

Output:
```
q7 results:
    q7 - 1 result:
        ❌ Test case failed
        Trying:
            points: 2
        Expecting nothing
        ok
        Trying:
            np.testing.assert_equal(get_3_nearest((3, 8), [(9, 3), (8, 5), (7, 6)]), [(7, 6), (8, 5), (9, 3)])
        Expecting nothing
        **********************************************************************
        Line 2, in q7 0
        Failed example:
            np.testing.assert_equal(get_3_nearest((3, 8), [(9, 3), (8, 5), (7, 6)]), [(7, 6), (8, 5), (9, 3)])
        Exception raised:
            Traceback (most recent call last):
              File "/home/kipkoech/anaconda3/lib/python3.11/doctest.py", line 1351, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q7 0[1]>", line 1, in <module>
                np.testing.assert_equal(get_3_nearest((3, 8), [(9, 3), (8, 5), (7, 6)]), [(7, 6), (8, 5), (9, 3)])
                                        ^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^
```

**Question 8:**  Write a function `diagonal_vector(M)` that returns a **numpy** array of the list of **absolute** values of the main diagonal entries in the matrix $M$

```python
def diagonal_vector(M):
  """
  This function returns a numpy array of the absolute values of the main diagonal entries in the matrix M.

  Args:
    M: A 2D numpy array representing a matrix.

  Returns:
    A numpy array containing the absolute values of the main diagonal entries of M.
  """

  # Get the diagonal entries of the matrix
  diagonal_entries = np.diagonal(M)

  # Calculate the absolute values of the diagonal entries
  absolute_values = np.abs(diagonal_entries)

  return absolute_values

# Example usage
M = np.array([[1, 2, 3], [4, 5, 6], [7, 8, 9]])

diagonal_vector_M = diagonal_vector(M)

print(f"Diagonal vector of M: {diagonal_vector_M}")
```

Output:
```
Diagonal vector of M: [1 5 9]
```

```python
grader.check("q8")
```

Output:
```
q8 results: All test cases passed!
```

**Question 9:**  Write a function `flatten_reverse_lists` that takes in a list of lists and outputs a **reverse** sorted list of elements of sublists of the input list (confusing right?) <br>
Example: given `flatten_reverse_lists([[2,13,44], [6,7]])` it should return `[2,6,7,13,44]`

```python
def flatten_reverse_lists(superlist):
  """
  This function takes in a list of lists and outputs a reverse sorted list of elements of sublists of the input list.

  Args:
      list_of_lists: A list of lists.

  Returns:
      A reverse sorted list of elements of sublists of the input list.
  """

  # Flatten the list of lists
  flattened_list = [item for sublist in superlist for item in sublist]

  # Sort the flattened list in reverse order
  flattened_list.sort(reverse=True)

  return flattened_list

# Example usage
superlist = [[2, 13, 44], [6, 7, 3 ], [ 1, 2, 3]]

flattened_reverse_list = flatten_reverse_lists(superlist)

print(f"Flattened reverse list: {flattened_reverse_list}")
```

Output:
```
Flattened reverse list: [44, 13, 7, 6, 3, 3, 2, 2, 1]
```

```python
grader.check("q9")
```

Output:
```
q9 results: All test cases passed!
```

**Question 9:** Create a DataFrame mirroring the table below and assign this to `data`. 

| flavor | scoops | price |
|-----|-----|-----|
| white chocolate | 1 | 2 |
| vanilla | 1 | 1.5 |
| dark chocolate | 2 | 3 |
| strawberry | 1 | 2 |
| strawberry | 3 | 4 |
| vanilla | 2 | 2 |
| mint | 1 | 4 |
| mint | 2 | 5 |
| white chocolate | 3 | 2 |
| dark chocolate | 3 | 3 |
| white chocolate | 2 | 2 |
| dark chocolate | 5 | 3 |


```python
import pandas as pd
data = [['white chokolate', 1, 2], ['vanilla', 1, 1.5], ['dark chocolate' ,2 ,3],['strawberry ',1,2],['strawberry', 3,4],['vanila',2,2],['mint',1,4],['mint',2,5],['white chokolate', 3,2],['dark chocolate', 3,3],['white chocolate', 2,2],['dark chocolate',3,5]] 
df = pd.DataFrame(data, columns= ['flavor','scoops', 'price'])
df.reset_index(drop=False)
# df= pd.DataFrame(data)
df

```

Output:
```
flavor  scoops  price
0   white chokolate       1    2.0
1           vanilla       1    1.5
2    dark chocolate       2    3.0
3       strawberry        1    2.0
4        strawberry       3    4.0
5            vanila       2    2.0
6              mint       1    4.0
7              mint       2    5.0
8   white chokolate       3    2.0
9    dark chocolate       3    3.0
10  white chocolate       2    2.0
11   dark chocolate       3    5.0
```

**Question 10:** Do the following to the dataframe:
* Create a new collumn ``total_price`` whose value is equal to $scoops * price$*
* Write a function ``groupStatistics(data, groupValue)``. Internally, this function groups ``data``  by ``flavor`` and then returns statistics of a given grouped item ``groupValue`` indexed on the ``total_price`` columns. The statistics is a numpy array contains ``[mean, media, min, max, std]`` of the ``total_price`` column. The ``std`` should be rounded to 2 **decimal places**



```python
import pandas as pd
import numpy as np

def groupStatistics(data, groupValue):

  # Group data by flavor
    grouped_data = data.groupby(groupValue)

  # Calculate statistics for each group
    statistics = grouped_data["total_price"].agg(["mean", "median", "min", "max", "std"]).round(2).reset_index()

  # Convert the statistics to a numpy array
    statistics_array = statistics.to_numpy()

    return statistics_array


```

Output:
```
[['dark chocolate' 10.0 9.0 6.0 15.0 4.58]
 ['mint' 7.0 7.0 4.0 10.0 4.24]
 ['strawberry' 7.0 7.0 2.0 12.0 7.07]
 ['vanila' 4.0 4.0 4.0 4.0 nan]
 ['vanilla' 1.5 1.5 1.5 1.5 nan]
 ['white chocolate' 4.0 4.0 4.0 4.0 nan]
 ['white chokolate' 4.0 4.0 2.0 6.0 2.83]]
```

```python
grader.check("q11")
```

Output:
```
q11 results:
    q11 - 1 result:
        ❌ Test case failed
        Trying:
            points: 1
        Expecting nothing
        ok
        Trying:
            np.testing.assert_array_equal(groupStatistics(data, 'strawberry'), [7.0, 7.0, 2.0, 7.07])
        Expecting nothing
        **********************************************************************
        Line 2, in q11 0
        Failed example:
            np.testing.assert_array_equal(groupStatistics(data, 'strawberry'), [7.0, 7.0, 2.0, 7.07])
        Exception raised:
            Traceback (most recent call last):
              File "/home/kipkoech/anaconda3/lib/python3.11/doctest.py", line 1351, in __run
                exec(compile(example.source, filename, "single",
              File "<doctest q11 0[1]>", line 1, in <module>
                np.testing.assert_array_equal(groupStatistics(data, 'strawberry'), [7.0, 7.0, 2.0, 7.07])
                                              ^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^^
           
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

Download the exported ZIP. Take note of the ZIP number and proceed to fill the summer school form

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Output:
```
<IPython.core.display.HTML object>
```

Tests for q1 (points: 4)
```
>>> points: 1
>>> 
>>> def test_isValid(isValid):
...     assert isValid('{[()]}}') == False
>>> test_isValid(isValid)
```
```
>>> points: 1
>>> 
>>> def test_isValid(isValid):
...     assert isValid('') == True
>>> test_isValid(isValid)
```
```
>>> points: 1
>>> 
>>> def test_isValid(isValid):
...     assert isValid('(){}[]') == True
>>> test_isValid(isValid)
```


Tests for q11 (points: 5)
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'strawberry'), [7.0, 7.0, 2.0, 7.07])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'dark chocolate'), [10.0, 9.0, 6.0, 4.58])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'white chocolate'), [4.0, 4.0, 2.0, 2.0])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'vanilla'), [2.75, 2.75, 1.5, 1.77])
```
```
>>> points: 1
>>> np.testing.assert_array_equal(groupStatistics(data, 'mint'), [7.0, 7.0, 4.0, 4.24])
```


Tests for q2 (points: 4)
```
>>> points: 1
>>> 
>>> def test_oddFrequencyCounter(oddFrequencyCounter):
...     paragraph = 'DSA 2024 Nyeri'
...     assert oddFrequencyCounter(paragraph) == 10
```
```
>>> points: 1
>>> 
>>> def test_oddFrequencyCounter(oddFrequencyCounter):
...     paragraph = "The following is sample text I used to practice special characters using keybr.com: 112233445566778899 Saturn V rocket's first stage carries 203,400 gallons (770,000 liters) of kerosene fuel and 318,000 gallons (1.2 million liters) of liquid oxygen needed for combustion. At liftoff, the stage's five F-1 rocket engines ignite and produce 7.5 million pounds of thrust. To replace those goofy quantities with the far less retarded metric system (even though liters are considered part of the metric system they are the same as cubic deci-meters) you would say 770 cubic meters of kerosene {abbreviated as m3} and 1,204 m3 of liquid O2 [O2 is the symbol for oxygen]. We would also say it produced 33,600,000 newtons of force <abbreviated as N>. To add scientific notation {a way of writing numbers that allows you to write only as many digits of specificity as you would like} you can write 7.7 * 10 ^ 2 m3 of kerosene 1.204 * 10 ^ 3 m3 of O2 and 3.3 * 10 ^ 7 newtons."
...     assert oddFrequencyCounter(paragraph) == 40
```
```
>>> points: 1
>>> 
>>> def test_oddFrequencyCounter(oddFrequencyCounter):
...     paragraph = 'Data Science Africa in collaboration with The Swiss Federal Department of Foreign Affairs (FDFA), ETH Zurich, EPFL, the European Laboratory for Learning and Intelligent Systems (ELLIS), the Swiss National Computing Centre (CSCS), and the LUMI consortium launched the International Computation and AI Network (ICAIN) at the World Economic Forum (WEF) 2024 in Davos. Its mission is to develop AI technologies that benefit society as a whole, as well as being accessible to all and sustainable, thereby helping to reduce global inequality.'
...     assert oddFrequencyCounter(paragraph) == 20
```


Tests for q3 (points: 4)
```
>>> points: 4
>>> theGen = odd_squares_sum(12)
>>> assert next(theGen) == 1
>>> assert next(theGen) == 9
>>> assert next(theGen) == 25
>>> assert next(theGen) == 49
>>> assert next(theGen) == 81
>>> assert next(theGen) == 121
>>> del theGen
```


Tests for q4 (points: 4)
```
>>> points: 2
>>> assert (oddSumList == np.array([1, 9, 25, 49, 81, 121, 169, 225, 289, 361])).all()
```
```
>>> points: 2
>>> assert isinstance(oddSumList, np.ndarray)
```


Tests for q5 (points: 2)
```
>>> points: 2
>>> assert mergedList == [(1, 1), (4, 9), (0, 25), (4, 49), (1, 81), (1, 121), (4, 169), (0, 225), (4, 289), (1, 361)]
```


Tests for q6 (points: 5)
```
>>> points: 1
>>> np.testing.assert_equal(greatest_common_divisor(10, 15), 5)
```
```
>>> points: 1
>>> np.testing.assert_equal(greatest_common_divisor(15, 19), 1)
```
```
>>> points: 1
>>> np.testing.assert_equal(greatest_common_divisor(100, 105), 5)
```


Tests for q7 (points: 8)
```
>>> points: 2
>>> np.testing.assert_equal(get_3_nearest((3, 8), [(9, 3), (8, 5), (7, 6)]), [(7, 6), (8, 5), (9, 3)])
```
```
>>> points: 2
>>> np.testing.assert_equal(get_3_nearest((12, 8), [(5, 9), (9, 1), (2, 4), (13, 9), (10, 12)]), [(13, 9), (10, 12), (5, 9)])
```
```
>>> points: 2
>>> np.testing.assert_equal(get_3_nearest((5, 8), [(5, 9), (9, 1), (2, 4), (13, 9), (10, 12)]), [(5, 9), (2, 4), (13, 9)])
```


Tests for q8 (points: 4)
```
>>> points: 2
>>> np.testing.assert_array_equal(diagonal_vector([[1, 2, 3, 4, 5, 6, 7, 8], [9, 10, 11, 12, 13, 14, 15, 16], [17, 18, 19, 20, 21, 22, 23, 24], [25, 26, 27, 28, 29, 30, 31, 32], [33, 34, 35, 36, 37, 38, 39, 40], [41, 42, 43, 44, 45, 46, 47, 48], [49, 50, 51, 52, 53, 54, 55, 56], [57, 58, 59, 60, 61, 62, 63, 64]]), [1, 10, 19, 28, 37, 46, 55, 64])
```
```
>>> points: 2
>>> assert isinstance(diagonal_vector([[1, 2], [3, 4]]), np.ndarray)
```


Tests for q9 (points: 3)
```
>>> points: 1
>>> np.testing.assert_equal(flatten_reverse_lists([[2, 13, 44], [6, 7]]), [44, 13, 7, 6, 2])
```
```
>>> points: 1
>>> np.testing.assert_equal(flatten_reverse_lists([[2], [61, 34, 5, 8, 9]]), [61, 34, 9, 8, 5, 2])
```
